evaluation_interval: 1m

rule_files:
  - prometheus.gitops_deployment_alerts.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'kube_deployment_status_replicas_ready{name="depl-001", namespace="gitops-service-argocd"}'
        values: '0x11'
      - series: 'kube_deployment_status_replicas_ready{name="depl-002", namespace="gitops-service-argocd"}'
        values: '0x11'

      - series: 'kube_deployment_spec_replicas{name="depl-001", namespace="gitops-service-argocd"}'
        values: '1x11'
      - series: 'kube_deployment_spec_replicas{name="depl-002", namespace="gitops-service-argocd"}'
        values: '1x11'

    alert_rule_test:
      - eval_time: 10m
        alertname: GitOpsDeploymentHealthErrors
        exp_alerts:
          - exp_labels:
              severity: warning
              slo: true
            exp_annotations:
              summary: >-
                 Less than 95% of GitOps Deployments are in Healthy state: 0%
              description: >-
                The percentage total of Argo CD Deployments that are in Healthy state is less than 95%.
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/gitops/deployments.md

  - interval: 1m
    input_series:
      - series: 'kube_deployment_status_replicas_ready{name="depl-001", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '0x11'
      - series: 'kube_deployment_status_replicas_ready{name="depl-002", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '0x11'

      - series: 'kube_deployment_spec_replicas{name="depl-001", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '1x11'
      - series: 'kube_deployment_spec_replicas{name="depl-002", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '1x11'

    alert_rule_test:
      - eval_time: 10m
        alertname: GitOpsDeploymentHealthErrors
        exp_alerts:
          - exp_labels:
              severity: warning
              slo: true
              source_cluster: stonesoupp01ue1
            exp_annotations:
              summary: >-
                 Less than 95% of GitOps Deployments are in Healthy state: 0%
              description: >-
                The percentage total of Argo CD Deployments that are in Healthy state is less than 95%.
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/gitops/deployments.md

  - interval: 1m
    input_series:
      # We don't want to be alerted in this situation (don't care situation)
      - series: 'kube_deployment_status_replicas_ready{name="depl-001", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '_ _ _ _ _ _ _ _ _ _ _'
      - series: 'kube_deployment_status_replicas_ready{name="depl-002", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '_ _ _ _ _ _ _ _ _ _ _'

      - series: 'kube_deployment_spec_replicas{name="depl-001", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '_ _ _ _ _ _ _ _ _ _ _'
      - series: 'kube_deployment_spec_replicas{name="depl-002", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '_ _ _ _ _ _ _ _ _ _ _'

    alert_rule_test:
      - eval_time: 10m
        alertname: GitOpsDeploymentHealthErrors

  - interval: 1m
    input_series:
      # One healthy deployment, but unhealthy. 50% health rate, so alert
      - series: 'kube_deployment_status_replicas_ready{name="depl-001", namespace="gitops-service-argocd"}'
        values: '1x11'
      - series: 'kube_deployment_status_replicas_ready{name="depl-002", namespace="gitops-service-argocd"}'
        values: '0x11'

      - series: 'kube_deployment_spec_replicas{name="depl-001", namespace="gitops-service-argocd"}'
        values: '1x11'
      - series: 'kube_deployment_spec_replicas{name="depl-002", namespace="gitops-service-argocd"}'
        values: '1x11'

    alert_rule_test:
      - eval_time: 10m
        alertname: GitOpsDeploymentHealthErrors
        exp_alerts:
          - exp_labels:
              severity: warning
              slo: true
            exp_annotations:
              summary: >-
                 Less than 95% of GitOps Deployments are in Healthy state: 50%
              description: >-
                The percentage total of Argo CD Deployments that are in Healthy state is less than 95%.
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/gitops/deployments.md

  - interval: 1m
    input_series:
      # Same series as above but specify source_cluster
      - series: 'kube_deployment_status_replicas_ready{name="depl-001", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '1x11'
      - series: 'kube_deployment_status_replicas_ready{name="depl-002", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '0x11'

      - series: 'kube_deployment_spec_replicas{name="depl-001", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '1x11'
      - series: 'kube_deployment_spec_replicas{name="depl-002", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '1x11'

    alert_rule_test:
      - eval_time: 10m
        alertname: GitOpsDeploymentHealthErrors
        exp_alerts:
          - exp_labels:
              severity: warning
              slo: true
              source_cluster: stonesoupp01ue1
            exp_annotations:
              summary: >-
                 Less than 95% of GitOps Deployments are in Healthy state: 50%
              description: >-
                The percentage total of Argo CD Deployments that are in Healthy state is less than 95%.
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/gitops/deployments.md

  - interval: 1m
    input_series:
      # healthy deployment for entire window period. Basic scenario. Should not alert
      - series: 'kube_deployment_status_replicas_ready{name="depl-001", namespace="gitops-service-argocd"}'
        values: '1x11'
      - series: 'kube_deployment_status_replicas_ready{name="depl-002", namespace="gitops-service-argocd"}'
        values: '1x11'

      - series: 'kube_deployment_spec_replicas{name="depl-001", namespace="gitops-service-argocd"}'
        values: '1x11'
      - series: 'kube_deployment_spec_replicas{name="depl-002", namespace="gitops-service-argocd"}'
        values: '1x11'

    alert_rule_test:
      - eval_time: 6m
        alertname: GitOpsDeploymentHealthErrors

  - interval: 1m
    input_series:
      # Same series as above but with source_cluster
      - series: 'kube_deployment_status_replicas_ready{name="depl-001", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '1x11'
      - series: 'kube_deployment_status_replicas_ready{name="depl-002", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '1x11'

      - series: 'kube_deployment_spec_replicas{name="depl-001", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '1x11'
      - series: 'kube_deployment_spec_replicas{name="depl-002", namespace="gitops-service-argocd", source_cluster="stonesoupp01ue1"}'
        values: '1x11'

    alert_rule_test:
      - eval_time: 6m
        alertname: GitOpsDeploymentHealthErrors

  - interval: 1m
    input_series:
      # One deployment is unhealthy at the eval time 14 minutes into the series. 50% healthy, so alert.
      - series: 'kube_deployment_status_replicas_ready{name="depl-001", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 _ _ _ _ _ _ _ _ _ _ _ 1 1 1 1 1'
      - series: 'kube_deployment_status_replicas_ready{name="depl-002", namespace="gitops-service-argocd"}'
        values: '_ _ _ _ _ _ _ _ _ 1 1 1 1 1 _ _ _ _ _ _'

      - series: 'kube_deployment_spec_replicas{name="depl-001", namespace="gitops-service-argocd"}'
        values: '_ 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'
      - series: 'kube_deployment_spec_replicas{name="depl-002", namespace="gitops-service-argocd"}'
        values: '_ _ _ _ 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'

    alert_rule_test:
      - eval_time: 14m
        alertname: GitOpsDeploymentHealthErrors
        exp_alerts:
          - exp_labels:
              severity: warning
              slo: true
            exp_annotations:
              summary: >-
                 Less than 95% of GitOps Deployments are in Healthy state: 50%
              description: >-
                The percentage total of Argo CD Deployments that are in Healthy state is less than 95%.
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/gitops/deployments.md

  - interval: 1m
    input_series:
      # One deployment is unhealthy for 5 minutes. However, at the eval time at 20m, the alert
      # should not be firing and it missed the earlier window where it should have fired. But since the recent
      # measurements meet the SLO and perhaps the health status is stabilized, it is acceptable.
      - series: 'kube_deployment_status_replicas_ready{name="depl-001", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'
      - series: 'kube_deployment_status_replicas_ready{name="depl-002", namespace="gitops-service-argocd"}'
        values: '_ _ _ _ _ _ _ _ _ _ _ _ _ _ _ _ _ _ _ 1'

      - series: 'kube_deployment_spec_replicas{name="depl-001", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'
      - series: 'kube_deployment_spec_replicas{name="depl-002", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'

    alert_rule_test:
      - eval_time: 19m
        alertname: GitOpsDeploymentHealthErrors

  - interval: 1m
    input_series:
      # 5 deployments. 4 are healthy, 1 unhealthy. At the eval time of 20m, the last deployment has been uhealthy
      # so the alert should fire. 4/5 = 0.8
      - series: 'kube_deployment_status_replicas_ready{name="depl-001", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'
      - series: 'kube_deployment_status_replicas_ready{name="depl-002", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'
      - series: 'kube_deployment_status_replicas_ready{name="depl-003", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'
      - series: 'kube_deployment_status_replicas_ready{name="depl-004", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'
      - series: 'kube_deployment_status_replicas_ready{name="depl-005", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 1 1 1 1 1 _ _ _ _ _ _ _ _ _ _ _'

      - series: 'kube_deployment_spec_replicas{name="depl-001", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'
      - series: 'kube_deployment_spec_replicas{name="depl-002", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'
      - series: 'kube_deployment_spec_replicas{name="depl-003", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'
      - series: 'kube_deployment_spec_replicas{name="depl-004", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'
      - series: 'kube_deployment_spec_replicas{name="depl-005", namespace="gitops-service-argocd"}'
        values: '1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1'

    alert_rule_test:
      - eval_time: 20m
        alertname: GitOpsDeploymentHealthErrors
        exp_alerts:
          - exp_labels:
              severity: warning
              slo: true
            exp_annotations:
              summary: >-
                 Less than 95% of GitOps Deployments are in Healthy state: 80%
              description: >-
                The percentage total of Argo CD Deployments that are in Healthy state is less than 95%.
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/gitops/deployments.md
